[1637811929] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3381 MB out of 3792 MB)
[1637812229] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637812290] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (289 MB out of 3792 MB)
[1637812350] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637812410] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (265 MB out of 3792 MB)
[1637812711] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637813012] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 16% free (606 MB out of 3792 MB)
[1637813072] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (274 MB out of 3792 MB)
[1637813132] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637813192] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (239 MB out of 3792 MB)
[1637813792] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 80% free (3007 MB out of 3792 MB)
[1637814093] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (291 MB out of 3792 MB)
[1637814153] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (233 MB out of 3792 MB)
[1637814213] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637814274] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637814574] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637814875] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637814935] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (249 MB out of 3792 MB)
[1637814995] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637815056] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637815656] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 55% free (2084 MB out of 3792 MB)
[1637815957] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637816018] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (257 MB out of 3792 MB)
[1637816079] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637816139] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637817339] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3384 MB out of 3792 MB)
[1637817640] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (299 MB out of 3792 MB)
[1637817700] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (306 MB out of 3792 MB)
[1637817760] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637817821] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (251 MB out of 3792 MB)
[1637818121] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637818421] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (294 MB out of 3792 MB)
[1637818481] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637818541] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (301 MB out of 3792 MB)
[1637818602] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637818902] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637819502] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (308 MB out of 3792 MB)
[1637819563] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637819623] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637819683] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (290 MB out of 3792 MB)
[1637819983] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637820283] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (297 MB out of 3792 MB)
[1637820344] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (267 MB out of 3792 MB)
[1637820404] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (283 MB out of 3792 MB)
[1637820465] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (287 MB out of 3792 MB)
[1637821065] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 73% free (2764 MB out of 3792 MB)
[1637821366] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637821427] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (255 MB out of 3792 MB)
[1637821487] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (292 MB out of 3792 MB)
[1637821547] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (236 MB out of 3792 MB)
[1637822147] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 24% free (875 MB out of 3792 MB)
[1637822448] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (235 MB out of 3792 MB)
[1637822508] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (259 MB out of 3792 MB)
[1637822569] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637822629] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (258 MB out of 3792 MB)
[1637822930] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 34% free (1273 MB out of 3792 MB)
[1637823232] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (246 MB out of 3792 MB)
[1637823292] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637823352] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637823413] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (238 MB out of 3792 MB)
[1637823713] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3368 MB out of 3792 MB)
[1637824013] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637824074] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637824134] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637824194] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (261 MB out of 3792 MB)
[1637827194] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637827495] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (275 MB out of 3792 MB)
[1637827556] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (284 MB out of 3792 MB)
[1637827616] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (246 MB out of 3792 MB)
[1637827676] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637828276] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 28% free (1046 MB out of 3792 MB)
[1637828577] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637828637] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (238 MB out of 3792 MB)
[1637828697] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637828758] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637829058] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3385 MB out of 3792 MB)
[1637829358] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (244 MB out of 3792 MB)
[1637829418] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637829478] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (289 MB out of 3792 MB)
[1637829539] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637829839] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3383 MB out of 3792 MB)
[1637830139] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (263 MB out of 3792 MB)
[1637830200] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637830260] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637830321] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637830921] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 62% free (2334 MB out of 3792 MB)
[1637831221] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (261 MB out of 3792 MB)
[1637831282] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (247 MB out of 3792 MB)
[1637831342] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (254 MB out of 3792 MB)
[1637831403] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (286 MB out of 3792 MB)
[1637832604] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637832904] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 24% free (876 MB out of 3792 MB)
[1637833204] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637833264] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637833324] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637833385] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637833686] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 80% free (3024 MB out of 3792 MB)
[1637833988] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637834048] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (480 MB out of 3792 MB)
[1637834108] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (462 MB out of 3792 MB)
[1637834169] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (440 MB out of 3792 MB)
[1637834470] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 3% free (84 MB out of 3792 MB)
[1637834770] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637835371] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 2% free (67 MB out of 3792 MB)
[1637835672] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637836272] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3561 MB out of 3792 MB)
[1637836573] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637836634] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (435 MB out of 3792 MB)
[1637836694] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (442 MB out of 3792 MB)
[1637836754] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (444 MB out of 3792 MB)
[1637837054] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637837355] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637837416] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (473 MB out of 3792 MB)
[1637837476] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (475 MB out of 3792 MB)
[1637837536] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637837837] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637838437] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (443 MB out of 3792 MB)
[1637838497] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637838558] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637838619] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (442 MB out of 3792 MB)
[1637838919] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3546 MB out of 3792 MB)
[1637839219] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637839279] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637839340] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (472 MB out of 3792 MB)
[1637839400] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (466 MB out of 3792 MB)
[1637840000] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 63% free (2379 MB out of 3792 MB)
[1637840301] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637840361] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (434 MB out of 3792 MB)
[1637840422] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637840482] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (444 MB out of 3792 MB)
[1637840783] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 81% free (3059 MB out of 3792 MB)
[1637841083] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (469 MB out of 3792 MB)
[1637841143] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (475 MB out of 3792 MB)
[1637841204] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (454 MB out of 3792 MB)
[1637841264] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637841564] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637841865] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637841926] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637841986] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (440 MB out of 3792 MB)
[1637842046] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (446 MB out of 3792 MB)
[1637842346] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637842946] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637843007] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (479 MB out of 3792 MB)
[1637843068] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (474 MB out of 3792 MB)
[1637843128] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (460 MB out of 3792 MB)
[1637843428] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3557 MB out of 3792 MB)
[1637843728] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (455 MB out of 3792 MB)
[1637843789] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (459 MB out of 3792 MB)
[1637843850] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637843910] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (449 MB out of 3792 MB)
[1637844210] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637844511] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (468 MB out of 3792 MB)
[1637844571] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (473 MB out of 3792 MB)
[1637844631] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (452 MB out of 3792 MB)
[1637844692] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637844992] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637845592] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (480 MB out of 3792 MB)
[1637845653] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (436 MB out of 3792 MB)
[1637845713] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (468 MB out of 3792 MB)
[1637845774] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637846074] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637846376] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637846436] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (448 MB out of 3792 MB)
[1637846496] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (370 MB out of 3792 MB)
[1637846557] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (463 MB out of 3792 MB)
[1637847157] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3545 MB out of 3792 MB)
[1637847458] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637847518] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (476 MB out of 3792 MB)
[1637847579] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637847640] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637847940] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3557 MB out of 3792 MB)
[1637848240] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (470 MB out of 3792 MB)
[1637848301] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (481 MB out of 3792 MB)
[1637848361] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (467 MB out of 3792 MB)
[1637848422] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637848722] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637849023] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (481 MB out of 3792 MB)
[1637849084] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (468 MB out of 3792 MB)
[1637849144] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637849204] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (457 MB out of 3792 MB)
[1637849505] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637850106] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (467 MB out of 3792 MB)
[1637850166] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637850227] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637850287] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637850587] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637850887] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (408 MB out of 3792 MB)
[1637850947] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637851008] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637851068] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (439 MB out of 3792 MB)
[1637851368] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637851968] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (455 MB out of 3792 MB)
[1637852028] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637852088] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (454 MB out of 3792 MB)
[1637852148] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (440 MB out of 3792 MB)
[1637852448] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3557 MB out of 3792 MB)
[1637852749] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637852809] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (439 MB out of 3792 MB)
[1637852869] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637852929] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637853529] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 54% free (2015 MB out of 3792 MB)
[1637853830] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637853890] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637853950] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637854011] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (457 MB out of 3792 MB)
[1637854311] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 62% free (2326 MB out of 3792 MB)
[1637854612] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (454 MB out of 3792 MB)
[1637854672] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (454 MB out of 3792 MB)
[1637854732] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (461 MB out of 3792 MB)
[1637854792] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (473 MB out of 3792 MB)
[1637855092] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637855393] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637855453] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637855513] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (417 MB out of 3792 MB)
[1637855573] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637855873] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3544 MB out of 3792 MB)
[1637856473] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637856534] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (442 MB out of 3792 MB)
[1637856594] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637856654] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637857314] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637857375] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637857435] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (469 MB out of 3792 MB)
[1637857495] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (460 MB out of 3792 MB)
[1637857796] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637858096] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (465 MB out of 3792 MB)
[1637858156] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (466 MB out of 3792 MB)
[1637858216] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637858277] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (462 MB out of 3792 MB)
[1637858577] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637859178] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (434 MB out of 3792 MB)
[1637859238] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (442 MB out of 3792 MB)
[1637859298] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (457 MB out of 3792 MB)
[1637859359] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637860020] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (463 MB out of 3792 MB)
[1637860081] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (467 MB out of 3792 MB)
[1637860141] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637860201] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (443 MB out of 3792 MB)
[1637860501] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637860801] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (467 MB out of 3792 MB)
[1637860861] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (398 MB out of 3792 MB)
[1637860921] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637860982] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637861583] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 74% free (2769 MB out of 3792 MB)
[1637861883] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637861943] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (423 MB out of 3792 MB)
[1637862003] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (412 MB out of 3792 MB)
[1637862064] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (409 MB out of 3792 MB)
[1637862364] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3559 MB out of 3792 MB)
[1637862664] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637862725] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637862786] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637862846] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (475 MB out of 3792 MB)
[1637863146] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3532 MB out of 3792 MB)
[1637863447] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (402 MB out of 3792 MB)
[1637863507] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (479 MB out of 3792 MB)
[1637863568] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (449 MB out of 3792 MB)
[1637863629] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (461 MB out of 3792 MB)
[1637863929] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637864530] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (444 MB out of 3792 MB)
[1637864591] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637864651] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (473 MB out of 3792 MB)
[1637864711] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (446 MB out of 3792 MB)
[1637865012] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3548 MB out of 3792 MB)
[1637865312] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (398 MB out of 3792 MB)
[1637865373] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (472 MB out of 3792 MB)
[1637865433] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (462 MB out of 3792 MB)
[1637865493] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (428 MB out of 3792 MB)
[1637865793] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3530 MB out of 3792 MB)
[1637866094] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 16% free (574 MB out of 3792 MB)
[1637866155] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (463 MB out of 3792 MB)
[1637866216] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637866276] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (413 MB out of 3792 MB)
[1637866576] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637867177] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (478 MB out of 3792 MB)
[1637867237] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (476 MB out of 3792 MB)
[1637867297] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637867358] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (422 MB out of 3792 MB)
[1637867658] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3530 MB out of 3792 MB)
[1637867959] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (406 MB out of 3792 MB)
[1637868019] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (448 MB out of 3792 MB)
[1637868080] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637868140] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (416 MB out of 3792 MB)
[1637868440] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637869041] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (428 MB out of 3792 MB)
[1637869102] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (423 MB out of 3792 MB)
[1637869163] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (436 MB out of 3792 MB)
[1637869223] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637869523] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3526 MB out of 3792 MB)
[1637869823] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637869884] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637869944] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637870004] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (434 MB out of 3792 MB)
[1637870304] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637870605] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (430 MB out of 3792 MB)
[1637870666] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637870726] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637870787] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (434 MB out of 3792 MB)
[1637871087] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3509 MB out of 3792 MB)
[1637871688] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (400 MB out of 3792 MB)
[1637871748] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637871808] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637871869] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (380 MB out of 3792 MB)
[1637873069] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637873370] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (430 MB out of 3792 MB)
[1637873430] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (397 MB out of 3792 MB)
[1637873491] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (411 MB out of 3792 MB)
[1637873551] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637873851] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637874451] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (435 MB out of 3792 MB)
[1637874512] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (416 MB out of 3792 MB)
[1637874572] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (397 MB out of 3792 MB)
[1637874632] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637874933] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3534 MB out of 3792 MB)
[1637875233] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637875293] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (398 MB out of 3792 MB)
[1637875353] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (437 MB out of 3792 MB)
[1637875414] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (430 MB out of 3792 MB)
[1637875714] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637876017] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637876077] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (396 MB out of 3792 MB)
[1637876137] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637876198] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (424 MB out of 3792 MB)
[1637876798] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 75% free (2832 MB out of 3792 MB)
[1637877098] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637877159] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637877219] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (427 MB out of 3792 MB)
[1637877280] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (371 MB out of 3792 MB)
[1637877580] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637877881] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (394 MB out of 3792 MB)
[1637877941] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637878001] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (446 MB out of 3792 MB)
[1637878061] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637878362] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637878963] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (401 MB out of 3792 MB)
[1637879023] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637879083] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637879143] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (400 MB out of 3792 MB)
[1637879444] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3536 MB out of 3792 MB)
[1637879744] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (411 MB out of 3792 MB)
[1637879804] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637879865] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (422 MB out of 3792 MB)
[1637879925] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (402 MB out of 3792 MB)
[1637880225] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637880526] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637880586] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (426 MB out of 3792 MB)
[1637880647] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637880707] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637881007] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637881607] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (388 MB out of 3792 MB)
[1637881668] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (424 MB out of 3792 MB)
[1637881729] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637881789] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637882089] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637882389] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (418 MB out of 3792 MB)
[1637882450] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (413 MB out of 3792 MB)
[1637882510] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637882570] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637883171] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 62% free (2323 MB out of 3792 MB)
[1637883472] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (448 MB out of 3792 MB)
[1637883532] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (440 MB out of 3792 MB)
[1637883593] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (413 MB out of 3792 MB)
[1637883653] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637883953] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637884253] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (435 MB out of 3792 MB)
[1637884313] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (417 MB out of 3792 MB)
[1637884374] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637884434] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637884734] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637885035] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (404 MB out of 3792 MB)
[1637885095] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637885155] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637885216] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (417 MB out of 3792 MB)
[1637885516] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637886117] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (437 MB out of 3792 MB)
[1637886177] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637886238] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (439 MB out of 3792 MB)
[1637886298] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (422 MB out of 3792 MB)
[1637886598] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637886898] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637886959] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637887019] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637887080] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637887380] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637887980] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (443 MB out of 3792 MB)
[1637888041] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (463 MB out of 3792 MB)
[1637888101] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637888161] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 93% free (3515 MB out of 3792 MB)
[1637888823] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (452 MB out of 3792 MB)
[1637888884] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637888944] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637889004] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (404 MB out of 3792 MB)
[1637889605] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 22% free (829 MB out of 3792 MB)
[1637889905] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637889965] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637890026] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (450 MB out of 3792 MB)
[1637890086] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637890387] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 72% free (2714 MB out of 3792 MB)
[1637890687] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637890748] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (444 MB out of 3792 MB)
[1637890808] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637890868] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (424 MB out of 3792 MB)
[1637891169] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 1% free (2 MB out of 3792 MB)
[1637891469] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (409 MB out of 3792 MB)
[1637892070] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637892370] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (446 MB out of 3792 MB)
[1637892430] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (411 MB out of 3792 MB)
[1637892490] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (373 MB out of 3792 MB)
[1637892551] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637893153] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 46% free (1737 MB out of 3792 MB)
[1637893455] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637893515] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637893576] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (436 MB out of 3792 MB)
[1637893637] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637893937] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3533 MB out of 3792 MB)
[1637894238] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (416 MB out of 3792 MB)
[1637894298] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637894358] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (402 MB out of 3792 MB)
[1637894419] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (406 MB out of 3792 MB)
[1637894719] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (377 MB out of 3792 MB)
[1637895019] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (412 MB out of 3792 MB)