[1637762081] SERVICE ALERT: server-173;Current Load;OK;HARD;4;OK - load average: 1.98, 1.79, 1.52
[1637762090] SERVICE ALERT: server-173;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1637762120] SERVICE ALERT: server-173;Root / Partition;OK;HARD;4;DISK OK - free space: / 9744 MB (88% inode=94%):
[1637762131] SERVICE ALERT: server-173;Total Processes;OK;HARD;4;PROCS OK: 106 processes
[1637762274] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3398 MB out of 3792 MB)
[1637762420] SERVICE ALERT: server-113;Root / Partition;OK;HARD;4;DISK OK - free space: / 23242 MB (93% inode=95%):
[1637762420] SERVICE ALERT: server-113;Total Processes;OK;HARD;4;PROCS OK: 141 processes
[1637762499] SERVICE ALERT: server-113;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1637762642] SERVICE ALERT: server-113;Current Load;OK;HARD;4;OK - load average: 0.02, 0.05, 0.05
[1637762666] SERVICE ALERT: server-113;SWAP Uasge;OK;HARD;4;SWAP OK - 100% free (3344 MB out of 3347 MB)
[1637762874] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (268 MB out of 3792 MB)
[1637762878] SERVICE ALERT: server-164;Current Load;OK;HARD;4;OK - load average: 0.00, 0.00, 0.00
[1637762935] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637762965] SERVICE ALERT: server-164;SWAP Uasge;OK;HARD;4;SWAP OK - 100% free (3792 MB out of 3792 MB)
[1637762995] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (246 MB out of 3792 MB)
[1637763001] SERVICE ALERT: server-164;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1637763008] SERVICE ALERT: server-164;Root / Partition;OK;HARD;4;DISK OK - free space: / 9742 MB (88% inode=94%):
[1637763016] SERVICE ALERT: server-164;Total Processes;OK;HARD;4;PROCS OK: 92 processes
[1637763033] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637763333] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3384 MB out of 3792 MB)
[1637763633] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637763694] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (253 MB out of 3792 MB)
[1637763754] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637763814] SERVICE FLAPPING ALERT: server-173;SWAP Uasge;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1637763814] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (283 MB out of 3792 MB)
[1637764115] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637764416] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637764476] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637764536] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (287 MB out of 3792 MB)
[1637764597] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637764897] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637765497] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (285 MB out of 3792 MB)
[1637765557] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (288 MB out of 3792 MB)
[1637765618] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637765678] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (280 MB out of 3792 MB)
[1637765978] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637766279] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637766339] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (283 MB out of 3792 MB)
[1637766399] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (288 MB out of 3792 MB)
[1637766460] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (250 MB out of 3792 MB)
[1637766760] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3364 MB out of 3792 MB)
[1637767361] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (299 MB out of 3792 MB)
[1637767421] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (244 MB out of 3792 MB)
[1637767481] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637767542] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637767842] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3385 MB out of 3792 MB)
[1637768142] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (305 MB out of 3792 MB)
[1637768202] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637768262] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637768323] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (288 MB out of 3792 MB)
[1637768773] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3377 MB out of 3792 MB)
[1637769073] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (293 MB out of 3792 MB)
[1637769134] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (256 MB out of 3792 MB)
[1637769194] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (289 MB out of 3792 MB)
[1637769255] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637769555] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3372 MB out of 3792 MB)
[1637769855] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637769916] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (274 MB out of 3792 MB)
[1637769976] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (238 MB out of 3792 MB)
[1637770036] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (274 MB out of 3792 MB)
[1637770337] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3357 MB out of 3792 MB)
[1637770938] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (289 MB out of 3792 MB)
[1637770999] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (253 MB out of 3792 MB)
[1637771059] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637771119] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637771419] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3380 MB out of 3792 MB)
[1637771720] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637771780] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (293 MB out of 3792 MB)
[1637771840] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637771901] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (267 MB out of 3792 MB)
[1637772201] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637772502] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637772562] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637772623] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637772684] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (296 MB out of 3792 MB)
[1637772984] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637773584] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (245 MB out of 3792 MB)
[1637773645] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637773705] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (274 MB out of 3792 MB)
[1637773765] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (284 MB out of 3792 MB)
[1637774066] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637774366] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637774427] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (284 MB out of 3792 MB)
[1637774488] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (288 MB out of 3792 MB)
[1637774548] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (275 MB out of 3792 MB)
[1637774848] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637775449] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (292 MB out of 3792 MB)
[1637775510] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (244 MB out of 3792 MB)
[1637775571] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (253 MB out of 3792 MB)
[1637775631] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (285 MB out of 3792 MB)
[1637775931] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3383 MB out of 3792 MB)
[1637776231] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637776292] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (307 MB out of 3792 MB)
[1637776352] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (290 MB out of 3792 MB)
[1637776413] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637776713] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637777014] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637777074] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (302 MB out of 3792 MB)
[1637777134] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (261 MB out of 3792 MB)
[1637777195] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637777495] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3364 MB out of 3792 MB)
[1637778095] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637778155] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (249 MB out of 3792 MB)
[1637778216] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (284 MB out of 3792 MB)
[1637778276] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (273 MB out of 3792 MB)
[1637778576] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637778877] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (305 MB out of 3792 MB)
[1637778938] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637778998] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (233 MB out of 3792 MB)
[1637779059] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (275 MB out of 3792 MB)
[1637779359] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637779961] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637780021] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637780081] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (284 MB out of 3792 MB)
[1637780142] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637780442] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3369 MB out of 3792 MB)
[1637780742] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637780802] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (231 MB out of 3792 MB)
[1637780862] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (261 MB out of 3792 MB)
[1637780923] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (265 MB out of 3792 MB)
[1637782123] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637782424] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637782484] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (261 MB out of 3792 MB)
[1637782544] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (259 MB out of 3792 MB)
[1637782605] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637782905] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637783505] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 6% free (207 MB out of 3792 MB)
[1637783566] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637783627] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637783687] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (247 MB out of 3792 MB)
[1637783987] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637784287] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637784347] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637784407] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (285 MB out of 3792 MB)
[1637784468] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (290 MB out of 3792 MB)
[1637784768] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637785369] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637785430] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637785490] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (251 MB out of 3792 MB)
[1637785550] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (245 MB out of 3792 MB)
[1637785850] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3385 MB out of 3792 MB)
[1637786150] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637786211] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (242 MB out of 3792 MB)
[1637786272] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (244 MB out of 3792 MB)
[1637786332] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (236 MB out of 3792 MB)
[1637786633] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3380 MB out of 3792 MB)
[1637786933] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637786993] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637787053] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (285 MB out of 3792 MB)
[1637787113] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (257 MB out of 3792 MB)
[1637787413] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637788014] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637788074] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (262 MB out of 3792 MB)
[1637788135] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (264 MB out of 3792 MB)
[1637788195] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (241 MB out of 3792 MB)
[1637791195] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637791495] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637791556] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637791616] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (249 MB out of 3792 MB)
[1637791676] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (276 MB out of 3792 MB)
[1637792278] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 49% free (1850 MB out of 3792 MB)
[1637792579] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (302 MB out of 3792 MB)
[1637792639] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (290 MB out of 3792 MB)
[1637792700] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637792760] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637793060] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3384 MB out of 3792 MB)
[1637793360] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (263 MB out of 3792 MB)
[1637793420] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (239 MB out of 3792 MB)
[1637793481] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637793541] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (251 MB out of 3792 MB)
[1637794741] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637795043] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (293 MB out of 3792 MB)
[1637795103] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (263 MB out of 3792 MB)
[1637795163] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637795224] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637795824] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3344 MB out of 3792 MB)
[1637796124] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637796185] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 6% free (210 MB out of 3792 MB)
[1637796245] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (253 MB out of 3792 MB)
[1637796305] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (245 MB out of 3792 MB)
[1637797505] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3368 MB out of 3792 MB)
[1637797805] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (251 MB out of 3792 MB)
[1637797866] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (290 MB out of 3792 MB)
[1637797926] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637797986] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637798287] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637798587] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 18% free (647 MB out of 3792 MB)
[1637798648] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (253 MB out of 3792 MB)
[1637798709] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (266 MB out of 3792 MB)
[1637798769] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (252 MB out of 3792 MB)
[1637799369] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3387 MB out of 3792 MB)
[1637799669] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (259 MB out of 3792 MB)
[1637799730] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637799790] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (294 MB out of 3792 MB)
[1637799850] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (250 MB out of 3792 MB)
[1637800150] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637800451] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (237 MB out of 3792 MB)
[1637800511] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (256 MB out of 3792 MB)
[1637800572] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (281 MB out of 3792 MB)
[1637800632] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (269 MB out of 3792 MB)
[1637800932] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3363 MB out of 3792 MB)
[1637801533] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (268 MB out of 3792 MB)
[1637801593] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637801653] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (243 MB out of 3792 MB)
[1637801713] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (278 MB out of 3792 MB)
[1637802013] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3375 MB out of 3792 MB)
[1637802314] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (283 MB out of 3792 MB)
[1637802374] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637802435] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (286 MB out of 3792 MB)
[1637802495] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (303 MB out of 3792 MB)
[1637803097] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 71% free (2669 MB out of 3792 MB)
[1637803398] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (280 MB out of 3792 MB)
[1637803458] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (279 MB out of 3792 MB)
[1637803518] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (263 MB out of 3792 MB)
[1637803579] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (289 MB out of 3792 MB)
[1637803879] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 89% free (3362 MB out of 3792 MB)
[1637804180] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (272 MB out of 3792 MB)
[1637804240] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (260 MB out of 3792 MB)
[1637804301] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (270 MB out of 3792 MB)
[1637804361] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (251 MB out of 3792 MB)
[1637804962] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (458 MB out of 3792 MB)
[1637805263] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (268 MB out of 3792 MB)
[1637806765] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 31% free (1152 MB out of 3792 MB)
[1637807066] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (277 MB out of 3792 MB)
[1637807126] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 5% free (183 MB out of 3792 MB)
[1637807186] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (237 MB out of 3792 MB)
[1637807247] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (267 MB out of 3792 MB)
[1637807547] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 84% free (3161 MB out of 3792 MB)
[1637807848] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 7% free (265 MB out of 3792 MB)
[1637807908] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (271 MB out of 3792 MB)
[1637807968] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (248 MB out of 3792 MB)
[1637808029] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (259 MB out of 3792 MB)