[1637895620] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (322 MB out of 3792 MB)
[1637895920] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637896520] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (368 MB out of 3792 MB)
[1637896820] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (382 MB out of 3792 MB)
[1637897420] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3530 MB out of 3792 MB)
[1637897721] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (434 MB out of 3792 MB)
[1637897781] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637897841] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (395 MB out of 3792 MB)
[1637897902] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (436 MB out of 3792 MB)
[1637898505] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 59% free (2210 MB out of 3792 MB)
[1637898805] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637898865] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637898926] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (399 MB out of 3792 MB)
[1637898986] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (399 MB out of 3792 MB)
[1637899286] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3538 MB out of 3792 MB)
[1637899587] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (440 MB out of 3792 MB)
[1637899647] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637899708] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (409 MB out of 3792 MB)
[1637899768] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (426 MB out of 3792 MB)
[1637900068] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637900369] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637900429] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (413 MB out of 3792 MB)
[1637900489] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (427 MB out of 3792 MB)
[1637900550] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637901151] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 45% free (1700 MB out of 3792 MB)
[1637901452] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (451 MB out of 3792 MB)
[1637901512] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (408 MB out of 3792 MB)
[1637901574] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637901634] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (417 MB out of 3792 MB)
[1637901934] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3534 MB out of 3792 MB)
[1637902234] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (417 MB out of 3792 MB)
[1637902295] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637902355] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (405 MB out of 3792 MB)
[1637902415] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (457 MB out of 3792 MB)
[1637903016] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 58% free (2194 MB out of 3792 MB)
[1637903316] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (438 MB out of 3792 MB)
[1637903377] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (430 MB out of 3792 MB)
[1637903437] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (407 MB out of 3792 MB)
[1637903497] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (422 MB out of 3792 MB)
[1637903797] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 68% free (2558 MB out of 3792 MB)
[1637904098] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (423 MB out of 3792 MB)
[1637904158] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (412 MB out of 3792 MB)
[1637904219] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (393 MB out of 3792 MB)
[1637904280] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637904580] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637904880] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637904941] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (424 MB out of 3792 MB)
[1637905001] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (411 MB out of 3792 MB)
[1637905061] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (429 MB out of 3792 MB)
[1637905361] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3514 MB out of 3792 MB)
[1637905961] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637906022] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637906082] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637906142] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637906443] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 7% free (249 MB out of 3792 MB)
[1637906743] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (383 MB out of 3792 MB)
[1637907343] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3528 MB out of 3792 MB)
[1637907643] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (381 MB out of 3792 MB)
[1637907704] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (400 MB out of 3792 MB)
[1637907766] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637907826] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (387 MB out of 3792 MB)
[1637908428] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 59% free (2200 MB out of 3792 MB)
[1637908730] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (412 MB out of 3792 MB)
[1637908790] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637908850] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (413 MB out of 3792 MB)
[1637908911] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (427 MB out of 3792 MB)
[1637909211] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3539 MB out of 3792 MB)
[1637909512] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637909572] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637909632] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (430 MB out of 3792 MB)
[1637909693] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637910293] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 21% free (758 MB out of 3792 MB)
[1637910594] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (408 MB out of 3792 MB)
[1637910654] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (394 MB out of 3792 MB)
[1637910714] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (416 MB out of 3792 MB)
[1637910774] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (415 MB out of 3792 MB)
[1637911075] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 72% free (2716 MB out of 3792 MB)
[1637911376] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (379 MB out of 3792 MB)
[1637911436] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (381 MB out of 3792 MB)
[1637911497] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (392 MB out of 3792 MB)
[1637911557] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (417 MB out of 3792 MB)
[1637911857] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 3% free (77 MB out of 3792 MB)
[1637912158] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637912458] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (357 MB out of 3792 MB)
[1637913058] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (423 MB out of 3792 MB)
[1637913659] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 3% free (109 MB out of 3792 MB)
[1637913959] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (405 MB out of 3792 MB)
[1637914559] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3530 MB out of 3792 MB)
[1637914860] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (375 MB out of 3792 MB)
[1637914920] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (443 MB out of 3792 MB)
[1637914981] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (395 MB out of 3792 MB)
[1637915041] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (416 MB out of 3792 MB)
[1637915341] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3508 MB out of 3792 MB)
[1637915642] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (447 MB out of 3792 MB)
[1637915702] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (418 MB out of 3792 MB)
[1637915762] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (426 MB out of 3792 MB)
[1637915823] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (412 MB out of 3792 MB)
[1637916424] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 62% free (2343 MB out of 3792 MB)
[1637916724] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (431 MB out of 3792 MB)
[1637916785] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637916845] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637916905] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (395 MB out of 3792 MB)
[1637917206] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 93% free (3508 MB out of 3792 MB)
[1637917506] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637917566] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (425 MB out of 3792 MB)
[1637917627] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (394 MB out of 3792 MB)
[1637917687] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (420 MB out of 3792 MB)
[1637918288] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 68% free (2549 MB out of 3792 MB)
[1637918589] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637918650] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (446 MB out of 3792 MB)
[1637918710] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (421 MB out of 3792 MB)
[1637918770] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (378 MB out of 3792 MB)
[1637919370] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (419 MB out of 3792 MB)
[1637919970] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 1% free (5 MB out of 3792 MB)
[1637920270] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (388 MB out of 3792 MB)
[1637920871] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 3% free (89 MB out of 3792 MB)
[1637921171] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (445 MB out of 3792 MB)
[1637921772] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 3792 MB)
[1637922072] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637922672] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3562 MB out of 3792 MB)
[1637922972] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637923033] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (481 MB out of 3792 MB)
[1637923093] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 12% free (443 MB out of 3792 MB)
[1637923154] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637923454] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637923755] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637923815] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (449 MB out of 3792 MB)
[1637923875] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (465 MB out of 3792 MB)
[1637923936] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (453 MB out of 3792 MB)
[1637924236] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637924836] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (418 MB out of 3792 MB)
[1637924897] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (436 MB out of 3792 MB)
[1637924957] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (410 MB out of 3792 MB)
[1637925017] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (462 MB out of 3792 MB)
[1637925318] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (341 MB out of 3792 MB)
[1637925618] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (464 MB out of 3792 MB)
[1637926218] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3553 MB out of 3792 MB)
[1637926518] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (441 MB out of 3792 MB)
[1637926579] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 13% free (465 MB out of 3792 MB)
[1637926639] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (476 MB out of 3792 MB)
[1637926700] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 12% free (452 MB out of 3792 MB)
[1637927000] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637927301] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 13% free (480 MB out of 3792 MB)
[1637927362] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 12% free (432 MB out of 3792 MB)
[1637927423] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (459 MB out of 3792 MB)
[1637927483] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 13% free (460 MB out of 3792 MB)
[1637927783] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 94% free (3537 MB out of 3792 MB)
[1637928383] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 12% free (433 MB out of 3792 MB)
[1637928443] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;2;SWAP WARNING - 11% free (414 MB out of 3792 MB)
[1637928504] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 13% free (456 MB out of 3792 MB)
[1637928565] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (408 MB out of 3792 MB)
[1637929765] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 95% free (3569 MB out of 3792 MB)
[1637930066] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637930126] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (315 MB out of 3792 MB)
[1637930187] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (356 MB out of 3792 MB)
[1637930248] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (311 MB out of 3792 MB)
[1637930548] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637931149] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637931209] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637931270] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (350 MB out of 3792 MB)
[1637931330] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (331 MB out of 3792 MB)
[1637931630] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3449 MB out of 3792 MB)
[1637931930] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637931990] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (352 MB out of 3792 MB)
[1637932051] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (339 MB out of 3792 MB)
[1637932111] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637932411] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637932712] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 15% free (537 MB out of 3792 MB)
[1637932772] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (267 MB out of 3792 MB)
[1637932832] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637932893] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (320 MB out of 3792 MB)
[1637933193] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637933793] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (326 MB out of 3792 MB)
[1637933854] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (328 MB out of 3792 MB)
[1637933914] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (323 MB out of 3792 MB)
[1637933975] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637934275] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637934576] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637934636] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637934697] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637934757] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (317 MB out of 3792 MB)
[1637935359] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 50% free (1859 MB out of 3792 MB)
[1637935659] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637935720] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (341 MB out of 3792 MB)
[1637935780] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (324 MB out of 3792 MB)
[1637935840] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (320 MB out of 3792 MB)
[1637937040] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3449 MB out of 3792 MB)
[1637937340] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (340 MB out of 3792 MB)
[1637937400] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (339 MB out of 3792 MB)
[1637937461] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (319 MB out of 3792 MB)
[1637937521] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (317 MB out of 3792 MB)
[1637938122] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 11% free (385 MB out of 3792 MB)
[1637938422] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (348 MB out of 3792 MB)
[1637938723] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637939023] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (325 MB out of 3792 MB)
[1637939083] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (267 MB out of 3792 MB)
[1637939143] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (307 MB out of 3792 MB)
[1637939204] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (352 MB out of 3792 MB)
[1637939504] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637940105] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (339 MB out of 3792 MB)
[1637940166] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (348 MB out of 3792 MB)
[1637940227] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (356 MB out of 3792 MB)
[1637940287] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (337 MB out of 3792 MB)
[1637940587] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637940887] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (353 MB out of 3792 MB)
[1637940947] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (316 MB out of 3792 MB)
[1637941007] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637941068] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637941668] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 60% free (2264 MB out of 3792 MB)
[1637941969] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (349 MB out of 3792 MB)
[1637942030] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (357 MB out of 3792 MB)
[1637942090] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (353 MB out of 3792 MB)
[1637942150] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637942450] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3450 MB out of 3792 MB)
[1637942750] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637942811] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (338 MB out of 3792 MB)
[1637942872] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (312 MB out of 3792 MB)
[1637942932] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (337 MB out of 3792 MB)
[1637943233] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637943534] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (325 MB out of 3792 MB)
[1637943594] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637943655] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (341 MB out of 3792 MB)
[1637943715] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637944015] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637944615] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (307 MB out of 3792 MB)
[1637944676] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637944736] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637944796] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (341 MB out of 3792 MB)
[1637945996] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637946296] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (339 MB out of 3792 MB)
[1637946356] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (340 MB out of 3792 MB)
[1637946417] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637946477] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (364 MB out of 3792 MB)
[1637946777] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637947378] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (355 MB out of 3792 MB)
[1637947438] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (368 MB out of 3792 MB)
[1637947499] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (333 MB out of 3792 MB)
[1637947559] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637948220] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (368 MB out of 3792 MB)
[1637948280] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (359 MB out of 3792 MB)
[1637948341] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (350 MB out of 3792 MB)
[1637948401] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (318 MB out of 3792 MB)
[1637948702] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637949002] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637949062] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637949122] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (337 MB out of 3792 MB)
[1637949182] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (332 MB out of 3792 MB)
[1637949783] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 54% free (2043 MB out of 3792 MB)
[1637950083] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637950144] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (309 MB out of 3792 MB)
[1637950204] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (321 MB out of 3792 MB)
[1637950264] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637950564] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3450 MB out of 3792 MB)
[1637950865] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (348 MB out of 3792 MB)
[1637950926] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (359 MB out of 3792 MB)
[1637950986] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (367 MB out of 3792 MB)
[1637951047] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (353 MB out of 3792 MB)
[1637951347] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3450 MB out of 3792 MB)
[1637951648] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (318 MB out of 3792 MB)
[1637951708] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637951768] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;3;SWAP WARNING - 11% free (381 MB out of 3792 MB)
[1637951829] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (325 MB out of 3792 MB)
[1637952129] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637952730] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637952791] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 7% free (248 MB out of 3792 MB)
[1637952851] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (292 MB out of 3792 MB)
[1637952911] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (358 MB out of 3792 MB)
[1637953211] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3434 MB out of 3792 MB)
[1637953511] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (313 MB out of 3792 MB)
[1637953572] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (326 MB out of 3792 MB)
[1637953632] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637953692] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637953993] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637954293] SERVICE ALERT: server-173;SWAP Uasge;WARNING;SOFT;1;SWAP WARNING - 16% free (603 MB out of 3792 MB)
[1637954354] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (308 MB out of 3792 MB)
[1637954414] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637954474] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (358 MB out of 3792 MB)
[1637954775] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637955375] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (352 MB out of 3792 MB)
[1637955436] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (324 MB out of 3792 MB)
[1637955496] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (350 MB out of 3792 MB)
[1637955556] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (311 MB out of 3792 MB)
[1637955856] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637956158] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637956218] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (318 MB out of 3792 MB)
[1637956279] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (324 MB out of 3792 MB)
[1637956339] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (358 MB out of 3792 MB)
[1637956940] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 59% free (2211 MB out of 3792 MB)
[1637957240] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637957301] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (335 MB out of 3792 MB)
[1637957361] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (293 MB out of 3792 MB)
[1637957421] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637957721] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3444 MB out of 3792 MB)
[1637958022] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (330 MB out of 3792 MB)
[1637958082] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (311 MB out of 3792 MB)
[1637958143] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (348 MB out of 3792 MB)
[1637958204] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (351 MB out of 3792 MB)
[1637958805] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 14% free (495 MB out of 3792 MB)
[1637959106] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637959708] SERVICE ALERT: server-173;SWAP Uasge;WARNING;HARD;4;SWAP WARNING - 15% free (552 MB out of 3792 MB)
[1637960009] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (321 MB out of 3792 MB)
[1637960609] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 58% free (2184 MB out of 3792 MB)
[1637960910] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (349 MB out of 3792 MB)
[1637960971] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637961031] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (316 MB out of 3792 MB)
[1637961092] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (326 MB out of 3792 MB)
[1637961392] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 90% free (3395 MB out of 3792 MB)
[1637961692] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (338 MB out of 3792 MB)
[1637961753] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637961813] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (332 MB out of 3792 MB)
[1637961874] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (328 MB out of 3792 MB)
[1637962174] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637962474] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637962534] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (324 MB out of 3792 MB)
[1637962595] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637962656] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (357 MB out of 3792 MB)
[1637962956] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637963557] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637963618] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (315 MB out of 3792 MB)
[1637963678] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637963738] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637964399] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (358 MB out of 3792 MB)
[1637964460] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (340 MB out of 3792 MB)
[1637964520] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637964581] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (306 MB out of 3792 MB)
[1637964881] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637965182] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (355 MB out of 3792 MB)
[1637965242] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (333 MB out of 3792 MB)
[1637965302] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (372 MB out of 3792 MB)
[1637965362] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637965663] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637966263] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637966324] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637966384] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637966444] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637967105] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (318 MB out of 3792 MB)
[1637967166] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637967227] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (280 MB out of 3792 MB)
[1637967287] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (325 MB out of 3792 MB)
[1637967587] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637967887] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (334 MB out of 3792 MB)
[1637967947] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (332 MB out of 3792 MB)
[1637968008] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 8% free (294 MB out of 3792 MB)
[1637968069] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (316 MB out of 3792 MB)
[1637968669] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 54% free (2027 MB out of 3792 MB)
[1637968969] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637969029] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (348 MB out of 3792 MB)
[1637969090] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637969150] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637969450] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3450 MB out of 3792 MB)
[1637969750] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (300 MB out of 3792 MB)
[1637969810] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (356 MB out of 3792 MB)
[1637969871] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637969931] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (343 MB out of 3792 MB)
[1637970231] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637970532] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (326 MB out of 3792 MB)
[1637970592] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (350 MB out of 3792 MB)
[1637970653] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (340 MB out of 3792 MB)
[1637970713] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (342 MB out of 3792 MB)
[1637971014] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637971614] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637971675] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (319 MB out of 3792 MB)
[1637971736] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (330 MB out of 3792 MB)
[1637971796] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (344 MB out of 3792 MB)
[1637972096] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3428 MB out of 3792 MB)
[1637972396] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (329 MB out of 3792 MB)
[1637972457] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (326 MB out of 3792 MB)
[1637972517] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (349 MB out of 3792 MB)
[1637972578] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (355 MB out of 3792 MB)
[1637972878] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637973478] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637973539] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (370 MB out of 3792 MB)
[1637973599] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (316 MB out of 3792 MB)
[1637973659] SERVICE ALERT: server-173;SWAP Uasge;OK;SOFT;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637974319] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (336 MB out of 3792 MB)
[1637974380] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (345 MB out of 3792 MB)
[1637974440] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (346 MB out of 3792 MB)
[1637974500] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (302 MB out of 3792 MB)
[1637974800] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637975101] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (308 MB out of 3792 MB)
[1637975161] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (339 MB out of 3792 MB)
[1637975222] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (350 MB out of 3792 MB)
[1637975284] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (361 MB out of 3792 MB)
[1637975887] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 39% free (1458 MB out of 3792 MB)
[1637976190] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (287 MB out of 3792 MB)
[1637976251] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 8% free (282 MB out of 3792 MB)
[1637976311] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (335 MB out of 3792 MB)
[1637976372] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (319 MB out of 3792 MB)
[1637976672] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3449 MB out of 3792 MB)
[1637976972] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (352 MB out of 3792 MB)
[1637977032] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (347 MB out of 3792 MB)
[1637977093] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (311 MB out of 3792 MB)
[1637977153] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (331 MB out of 3792 MB)
[1637977453] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 91% free (3427 MB out of 3792 MB)
[1637977754] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 9% free (331 MB out of 3792 MB)
[1637977815] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (376 MB out of 3792 MB)
[1637977876] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 10% free (358 MB out of 3792 MB)
[1637977937] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (342 MB out of 3792 MB)
[1637978538] SERVICE ALERT: server-173;SWAP Uasge;OK;HARD;4;SWAP OK - 60% free (2245 MB out of 3792 MB)
[1637978838] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;1;SWAP CRITICAL - 10% free (352 MB out of 3792 MB)
[1637978900] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (314 MB out of 3792 MB)
[1637978960] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;SOFT;3;SWAP CRITICAL - 9% free (314 MB out of 3792 MB)
[1637979021] SERVICE ALERT: server-173;SWAP Uasge;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (304 MB out of 3792 MB)